# Descriptive Stats

For the questions below, we're going to refer exclusively to an HHCAPS survey data set that is available in `/data/hhcaps.csv`.  Use whatever commands you want to calculate the information required to get to the answer.


Put your solution as the last line right before the test cell, using the form shown below for computing the number of rows in a dataframe.

```
answer = my_df.shape[0]
```

The assertions will often give you what the final answer should be, but you won't receive any points unless you compute the answer using code.  For instance, if you just typed `answer = 132` for the numer of rows instead of `answer = my_df.shape[0]`, you would not receive credit.


```python
%matplotlib inline
import pandas as pd

hhcaps = pd.read_csv('/data/hhcaps.csv')
```

## #01 - 

How many columns does this file contain?

```python
hhcaps.shape
```

Output:
```
(12000, 39)
```

<!-- grade_id: q01-solution, points: 0, grade: false, solution: true -->
```python
answer = hhcaps.shape[1]
## Solution goes here
```

<!-- grade_id: q01-test, points: 2, grade: true, solution: false -->
```python
assert(type(answer) == int)
assert(answer == 39)
```

## #02 - 

How many different values for State are there in this file?

```python
hhcaps.State.unique()
```

Output:
```
array(['AL', 'CO', 'AZ', 'FL', 'AK', 'GA', 'AR', 'HI', 'ID', 'CA', 'IL',
       'KY', 'LA', 'ME', 'MD', 'NY', 'MA', 'MO', 'NJ', 'NM', 'NC', 'OH',
       'PA', 'RI', 'SC', 'TX', 'CT', 'OK', 'VT', 'VA', 'DE', 'WA', 'DC',
       'WV', 'WI', 'IN', 'IA', 'KS', 'MI', 'MN', 'MS', 'NE', 'MT', 'NV',
       'NH', 'ND', 'OR', 'PR', 'SD', 'TN', 'UT', 'VI', 'WY', 'GU', 'MP'],
      dtype=object)
```

```python
hhcaps.State.nunique()
```

Output:
```
55
```

<!-- grade_id: q02-solution, points: 0, grade: false, solution: true -->
```python
answer = hhcaps.State.nunique()

## Solution goes here
```

<!-- grade_id: q02-test, points: 2, grade: true, solution: false -->
```python
assert(type(answer) == int)
assert(answer == 55)
```

```python

```

## #03 - 

Which of those State values has the highest frequency of occurence?

```python
hhcaps.State.value_counts()
```

Output:
```
TX    2518
CA    1319
FL    1069
OH     781
IL     735
MI     597
PA     425
MA     261
OK     256
VA     244
IN     229
MN     197
LA     195
CO     187
NC     174
AR     172
AZ     169
MO     169
IA     163
AL     152
NV     143
NY     141
TN     140
WI     119
KS     117
GA     106
KY     104
UT      98
CT      91
NM      77
NE      75
SC      70
WA      63
WV      61
OR      57
MD      53
MS      48
NJ      47
PR      45
ID      45
SD      33
NH      31
DC      28
RI      27
MT      27
DE      26
WY      26
ME      24
ND      18
HI      16
AK      14
VT      11
GU       4
MP       2
VI       1
Name: State, dtype: int64
```

<!-- grade_id: q03-solution, points: 0, grade: false, solution: true -->
```python
answer = hhcaps.State.value_counts().index[0]
## Solution goes here
```

<!-- grade_id: q03-test, points: 2, grade: true, solution: false -->
```python
assert(type(answer) == str)
assert(answer == 'TX')
```

## #04 - 

Which of those State values has the best average performance on the `Star Rating for health team communicated well with them` score?

```python
avg_srhtcw = hhcaps.groupby('State')['Star Rating for health team communicated well with them'].mean()
answer = avg_srhtcw[avg_srhtcw == avg_srhtcw.max()].index[0]

## Solution goes here

```

<!-- grade_id: q04-test, points: 2, grade: true, solution: false -->
```python
assert(type(answer) == str)
assert(answer == 'ME')
```

## #05 - 

What was the average score on `Star Rating for how patients rated overall care from agency` for providers listed as having a `Type of Ownership` of `Hospital Based Program`

<!-- grade_id: q05-solution, points: 0, grade: false, solution: true -->
```python
mask = hhcaps['Type of Ownership']=='Hospital Based Program'
filtered_hhcaps = hhcaps[mask]
answer = filtered_hhcaps['Star Rating for how patients rated overall care from agency'].mean()

## Solution goes here

```

<!-- grade_id: q05-test, points: 2, grade: true, solution: false -->
```python
import numpy
assert(type(answer) == float or type(answer) == numpy.float64)
assert(round(answer,5) == round(3.6925207756232687,5))

```

## #06 - 

There is one column in the file that has the same value on every row.  What is the name of that column?

```python
all_uniques = hhcaps.nunique()
answer = all_uniques[all_uniques==1].index[0]

## Solution goes here

```

<!-- grade_id: q06-test, points: 2, grade: true, solution: false -->
```python
assert(type(answer) == str)
assert(answer.lower() == 'Offers Nursing Care Services'.lower())

```

## #07 - 

There are six (6) columns in the file that indicate (True or False) if the provider offers certain services.  Those column names all start with `Offers...`.  How many facilities offer **everything** that the survey was interested in asking about?  That is, how many rows have True in all six of those columns?

```python
offers_cols = hhcaps.columns[hhcaps.columns.str.startswith('Offers')]

number_of_offers = hhcaps[offers_cols].sum(axis=1)

answer = int((number_of_offers == len(offers_cols)).sum())

## Solution goes here
```

<!-- grade_id: q07-test, points: 2, grade: true, solution: false -->
```python
assert(type(answer) == int)
assert(answer == 9029)

```

## #08 - 

The survey results report `Number of completed Surveys` as one of the metrics.  What is the `median` of the number of completed surveys per facility?

Note that some rows don't have a valid number in them for the number of completed surveys.


```python
mask = hhcaps['Number of completed Surveys']=='Not Available'

hhcaps.loc[mask,'Number of completed Surveys' ] = float('nan')
hhcaps['Number of completed Surveys'] = hhcaps['Number of completed Surveys'].astype(float)

hhcaps['Number of completed Surveys'].median(skipna=True)

answer = hhcaps['Number of completed Surveys'].median(skipna=True)

## Solution goes here
```

<!-- grade_id: q08-test, points: 2, grade: true, solution: false -->
```python
assert(type(answer) == float or type(answer) == numpy.float64)
assert(answer == 84)

```

## #09 - 

How many providers are there from St. Louis, Missouri?

*Note that the City columnd may have trailing spaces that need to be accomodated or cleaned up*

```python
hhcaps[(hhcaps['City'].str.strip() == 'SAINT LOUIS') & (hhcaps['State']=='MO')]['Provider Name'].nunique()
```

Output:
```
23
```

<!-- grade_id: q09-solution, points: 0, grade: false, solution: true -->
```python
answer = hhcaps[(hhcaps['City'].str.strip() == 'SAINT LOUIS') & (hhcaps['State']=='MO')]['Provider Name'].nunique()

## Solution goes here
```

<!-- grade_id: q09-test, points: 2, grade: true, solution: false -->
```python
assert(type(answer) == int)
assert(answer == 23)

```

# #10 - 

Which state got the highest percentage of 'top box' scorings, using the `Percent of patients who gave their home health agency a rating of 9 or 10 on a scale from 0 (lowest) to 10 (highest)` field.

Ignore those rows where the data is **Not Available**

```python
col = 'Percent of patients who gave their home health agency a rating of 9 or 10 on a scale from 0 (lowest) to 10 (highest)'

filtered = hhcaps[hhcaps[col]!='Not Available'].reset_index(drop=True)

filtered[ col] = filtered[col].astype(float)

filtered.groupby('State')[col].mean().sort_values(ascending=False).index[0]

answer = filtered.groupby('State')[col].mean().sort_values(ascending=False).index[0]
field = 'Percent of patients who gave their home health agency a rating of 9 or 10 on a scale from 0 (lowest) to 10 (highest)'

## Solution goes here

```

<!-- grade_id: q10-test, points: 2, grade: true, solution: false -->
```python
assert(type(answer) == str)
assert(answer == 'MS')
```

```python

```